# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
import pandas as pd
import os
# Since the first column of the dataset feature the unnamed column, I used index_col = 0 to take this 
# unnamed column to index
path = 'assets'
file = 'NISPUF17.csv'

cdc_2017_data = pd.read_csv(os.path.join(path, file), index_col = 0)
cdc_2017_data.head()
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  AGECPOXR  \
1   128521     12852     2         NaN  235.916956     1031  2017       NaN   
2    10741      1074     2         NaN  957.353840     1068  2017       NaN   
3   220011     22001     2         NaN  189.611299     1050  2017       NaN   
4    86131      8613     1  675.430817  333.447418     1040  2017       NaN   
5   227141     22714     1  482.617748  278.768063     1008  2017       NaN   

   HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  XVRCTY6  \
1         2       1  ...                        NaN      NaN      NaN   
2         2       1  ...                        NaN      NaN      NaN   
3         2       3  ...                        NaN      NaN      NaN   
4         2       1  ...                        NaN      NaN      NaN   
5         2       1  ...                        NaN      NaN      NaN   

   XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BREAK_I  
1      NaN      NaN      NaN    
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import os
# Since the first column of the dataset feature the unnamed column, I used index_col = 0 to take this 
# unnamed column to index
path = 'assets'
file = 'NISPUF17.csv'

def proportion_of_education():
    # your code goes here
    
    path = 'assets'
    file = 'NISPUF17.csv'
    cdc_2017_data = pd.read_csv(os.path.join(path, file), index_col = 0)
    edu_dict = {}
    edu_dict['less than high school'] = cdc_2017_data['EDUC1'][cdc_2017_data['EDUC1'] ==1].count() / cdc_2017_data['EDUC1'].count()
    edu_dict['high school'] = cdc_2017_data['EDUC1'][cdc_2017_data['EDUC1'] ==2].count() / cdc_2017_data['EDUC1'].count()
    edu_dict['more than high school but not college'] = cdc_2017_data['EDUC1'][cdc_2017_data['EDUC1'] ==3].count() / cdc_2017_data['EDUC1'].count()
    edu_dict['college'] = cdc_2017_data['EDUC1'][cdc_2017_data['EDUC1'] ==4].count() / cdc_2017_data['EDUC1'].count()
    
    return edu_dict
    
    # YOUR CODE HERE
    #raise NotImplementedError()
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
np.mean(cdc_2017_data['P_NUMFLU'][cdc_2017_data['CBF_01'] ==1])
```

Output:
```
1.8799187420058687
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import os
def average_influenza_doses():
    # YOUR CODE HERE
    path = 'assets'
    file = 'NISPUF17.csv'
    cdc_2017_data = pd.read_csv(os.path.join(path, file), index_col = 0)
    
    avg_flu_breastmilk = np.mean(cdc_2017_data['P_NUMFLU'][cdc_2017_data['CBF_01'] ==1])
    avg_flu_no_breastmilk = np.mean(cdc_2017_data['P_NUMFLU'][cdc_2017_data['CBF_01'] ==2])
    
    return (avg_flu_breastmilk, avg_flu_no_breastmilk)
    #raise NotImplementedError()
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.00773`.

```python
# Play around
cdc_2017_data = pd.read_csv('NISPUF17.csv', index_col = 0)
len(cdc_2017_data['HAD_CPOX'][(cdc_2017_data['HAD_CPOX'] == 77)])
```

Output:
```
105
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import os
def chickenpox_by_sex():
    # YOUR CODE HERE
    path = 'assets'
    file = 'NISPUF17.csv'
    cdc_2017_data = pd.read_csv(os.path.join(path, file), index_col = 0)
    ratio_cpox_dict = {}
    
    male_cpox = len(cdc_2017_data['HAD_CPOX'][(cdc_2017_data['SEX'] == 1) & (cdc_2017_data['P_NUMVRC'] >= 1) & (cdc_2017_data['HAD_CPOX'] == 1)])
    male_no_cpox = len(cdc_2017_data['HAD_CPOX'][(cdc_2017_data['SEX'] == 1) & (cdc_2017_data['P_NUMVRC'] >= 1) & (cdc_2017_data['HAD_CPOX'] == 2)])
    ratio_cpox_dict['male'] = male_cpox/male_no_cpox
    
    female_cpox = len(cdc_2017_data['HAD_CPOX'][(cdc_2017_data['SEX'] == 2) & (cdc_2017_data['P_NUMVRC'] >= 1) & (cdc_2017_data['HAD_CPOX'] == 1)])
    female_no_cpox = len(cdc_2017_data['HAD_CPOX'][(cdc_2017_data['SEX'] == 2) & (cdc_2017_data['P_NUMVRC'] >= 1) & (cdc_2017_data['HAD_CPOX'] == 2)])
    ratio_cpox_dict['female'] = female_cpox/female_no_cpox
    
    return ratio_cpox_dict
    #raise NotImplementedError()
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    import os 
    path = 'assets'
    file = 'NISPUF17.csv'
    cdc_2017_data = pd.read_csv(os.path.join(path, file), index_col = 0)
    cdc_2017_data = cdc_2017_data[["P_NUMVRC", "HAD_CPOX"]].dropna()
    cdc_2017_data = cdc_2017_data[cdc_2017_data["HAD_CPOX"] != 77]
    

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(cdc_2017_data["HAD_CPOX"], cdc_2017_data["P_NUMVRC"])
    
    # just return the correlation
    return corr
    #raise NotImplementedError()
corr_chickenpox()
```

Output:
```
0.07044873460147985
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```